## Students
Please fill in your names and S/U-numbers:
* Student 1 name, S/U-number:
* Student 2 name, S/U-number:
* Student 3 name, S/U-number:

# Statistical Machine Learning 2020
# Assignment 1
# Deadline: 7 October 2020

## Instructions
* You can __work in groups__ (= max 3 people). __Write the full name and S/U-number of all team members in the header above.__
* Make sure you __fill in any place that says__ `YOUR CODE HERE` or "YOUR ANSWER HERE" __including comments, derivations, explanations, graphs, etc.__ This means that the elements and/or intermediate steps required to derive the answer have to be in the report. (Answers like 'No' or 'x=27.2' by themselves are not sufficient, even when they are the result of running your code.) If an exercise requires coding, explain briefly what the code does (in comments). All figures should have titles (descriptions), axis labels, and legends (if applicable).
* Please do not add new cells, __write the answers only in the provided cells__. Before you turn this problem in, make sure everything runs as expected. First, *restart the kernel* (in the menubar, select Kernel$\rightarrow$Restart) and then *run all cells* (in the menubar, select Cell$\rightarrow$Run All). The assignment was written in (and we strongly recommend using) Python 3 by using the corresponding Python 3 kernel for Jupyter.
* The assignment includes certain cells that contain tests. Most of the tests are marked as *hidden* and are used for automatic grading. NB: These hidden tests do not provide any feedback! There are also a couple of tests / checks that are visible, which are meant to help you avoid basic coding errors.
* __Upload reports to Brightspace as a single .ipynb file containing the submitter's S/U-number: 'SML20_as01_&lt;submitter-S/U-number&gt;.ipynb'__, e.g., 'SML20_as01_S123456.ipynb'. For those working in groups, it is sufficient if one team member uploads the solutions.
* For any problems or questions, send us an email, or just ask. Email addresses: G.Bucur@cs.ru.nl, Yuliya.Shapovalova@ru.nl, and tomc@cs.ru.nl .

## Introduction
Assignment 1 consists of 3 exercises:
1. Polynomial curve fitting (50 points),
2. Gradient descent algorithm (25 points),
3. Probability theory (25 points).
The assignment also includes an optional exercise (worth 10 bonus points).

## Libraries
First, we import the basic libraries necessary to develop this assignment. Of course you are free to import further libraries, if required, in the allotted cells.

```python
# Necessary imports (for solutions)
import math
import numpy as np
import matplotlib.pyplot as plt
from collections import namedtuple

# Set fixed random seed for reproducibility
np.random.seed(2020)
```

## Exercise 1 (weight 50)
Consider once more the $M$-th order polynomial 
\begin{equation*}
y(x;\mathbf{w}) = w_0 + w_1 x + \ldots + w_M x^M  = \sum_{j=0}^M w_j x^j 
\label{yxw} \tag{1}
\end{equation*}

### Exercise 1.1
Create the function $f(x) = 1 + \sin(6(x - 2))$. Generate a data set $\mathcal{D}_{10}$ of 10 noisy observations of this function. Take the 10 inputs spaced uniformly in range $[0,1]$, and assume that the noise is Gaussian with mean 0 and standard deviation 0.3. $\mathcal{D}_{10}$ will be the training set. In a similar way, generate an additional test set $\mathcal{T}$ of 100 noisy observations over the same interval. Plot both the function and observations in $\mathcal{D}_{10}$ in a single graph (similar to Bishop, Fig.1.2).

<!-- grade_id: cell-b828ba4e85011481, points: 0, grade: false, solution: true -->
```python
def f(x):
    """
    This function computes (x)=1+sin(6(x−2))
    
    Parameters
    ----------
    x : float
        Input number.
    
    Returns
    -------
    float
        Result of the function.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7f16a8d5d13e2c74, points: 2, grade: true, solution: false -->
```python
"""
Basic check that function f is correct.
"""
assert f(2) == 1
```

<!-- grade_id: cell-26a157de86b14a7e, points: 0, grade: false, solution: true -->
```python
"""
Generate a data set of N_train noisy observations of the function f. Take the inputs spaced uniformly
in range [0,1], and add Gaussian noise with mean 0 and standard deviation 0.3.

Variable names
--------------
N_train : int
    number of training observations
X_train : array
    N_train x 1 vector of x-coordinates, uniformly distributed on [0, 1]
t_train : array
    N_train x 1 vector with corresponding t-values, adding Gaussian noise
D_train : matrix
    N_train x 2 matrix, the training data created from X_train and t_train

N_test : int
    number of data points for testing
X_test : array
    N_test x 1 vector of random x-coordinates taken form a uniform distribution
t_test : array
    N_test x 1 vector with corresponding t-values, adding Gaussian noise
D_test : matrix
    N_test x 2 matrix, the test data created from X_test and t_test
"""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6ed1d922cd47cc3a, points: 5, grade: true, solution: false -->
```python
"""
Hidden test for variables N_train, X_train, t_train.
"""
```

### Exercise 1.2
Create a function `polynomial_curve_fit(D,M)` that takes as input a data set $\mathcal{D}_{N}$, consisting of $N$ input/output-pairs $\{x_n,t_n\}$, and a parameter $M$, representing the order of the polynomial in \eqref{yxw}, and outputs a vector of weights $\mathbf{w} = [w_0, \dots, w_M]$ that minimizes the sum-of-squares error function
\begin{equation*} E(\mathbf{w}) = \frac{1}{2} \sum_{n=1}^N \{ y(x_{n} ; \mathbf{w}) - t_{n} \} ^2 \tag{2} \end{equation*}
Hint: use the results from the Tutorial Exercises (Week 1, Exercise 5), and `np.linalg.solve` to solve a linear system of equations.

<!-- grade_id: cell-b548cc4bd8339def, points: 0, grade: false, solution: true -->
```python
def polynomial_curve_fit(D, M):
    ''' This functions computes the value of a polynomial with weights w on data points x.
    
    Parameters
    ----------
    D : array
        Input dataset D.
    M : int
        The degree of the polynomial.
    
    Returns
    -------
    float
        Fitted weight vector w that minimizes the sum-of-squares function.
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a9ef1e5be2f54c31, points: 10, grade: true, solution: false -->
```python
"""
Hidden test for polynomial_curve_fit.
"""
```

### Exercise 1.3
For the given dataset $\mathcal{D}_{10}$, run the `polynomial_curve_fit(D,M)` function for $M = [0, \dots, 9]$,  and, 
* Plot for various orders $M$ (at least for $M=0, M=1, M=3, M=9$) the resulting polynomial, together with the function $f$ and observations $\mathcal{D}_{10}$ (similar to Bishop, Fig 1.4)
* For each order $M \in [0, \dots, 9]$,  compute the root-mean-square error
\begin{equation*} E_{\text{RMS}} = \sqrt{2 E(\mathbf{w^*})/N} \tag{3} \end{equation*}
of the corresponding polynomial, evaluated on both the training set $\mathcal{D}_{10}$ and the testset $\mathcal{T}$.  Plot both as a function of $M$ in a single graph. (see Bishop, Fig.1.5).

First define the `polynomial` function to help you with calculating the predictions of outputs for the training and test data given w.

<!-- grade_id: cell-f4213b5c3db7fe63, points: 0, grade: false, solution: true -->
```python
def polynomial(x, w):
    ''' This functions computes the value of a polynomial with weights w on data points x.
    
    Parameters
    ----------
    x : float
        Set of x-coordinates for which to evaluate the polynomial.
    w : float
        Input weight vector of size M+1 (for polynomial of degree M).
    
    Returns
    -------
    float
        Values of polynomial with weights w evaluated at x.
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ce23788082e731b4, points: 2, grade: true, solution: false -->
```python
"""
Test for polynomial.
"""
assert np.array_equal(polynomial(np.array([1, 2]), np.array([1, 2, 3])),np.array([ 6., 17.]))
```

Now with the help of `polynomial` calculate the predictions. Then calculate the root-mean-square-error and 
create plots for various orders of $M$.

<!-- grade_id: cell-acea12b9aae9969d, points: 12, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Exercise 1.4
Repeat this procedure for a data set $\mathcal{D}_{40}$ of  40 observations (with the same noise level) and compare with the previous result.

<!-- grade_id: cell-2be39b72eb2e1565, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d12e3459bbe09e9a, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 1.5
Modify the `polynomial_curve_fit(D,M)` function to include an additional penalty parameter $\lambda$, for a procedure that solves the minimization problem for a modified error function with quadratic regularizer (weight decay), given as
\begin{equation*}
\tilde{E} = E + \frac{\lambda}{2} \sum_{j=0}^M w_j^2. \label{regerr} \tag{4}
\end{equation*}
Verify that the regularizer drives the weights of high order terms in the polynomial to zero, and see if you can reproduce and explain the effect observed in Bishop, Fig.1.8. (note that the values here are computed for our data, so they are not identical to the ones in Bishop)

|$\ln\lambda=$|$-\infty$ |-18    |-9    |-4   |0    |
|-------------|----------|-------|------|-----|-----|
|$w_0^*$      |0.87      |0.89   |1.13  |1.59 |1.02 |
|$w_1^*$      |-166.39   |25.44  |8.97  |-0.20|-0.38|
|$w_2^*$      |4264.93   |-178.33|-29.89|-2.73|-0.41|
|$w_3^*$      |-39400.69 |530.85 |11.83 |-1.24|-0.25|
|$w_4^*$      |185746.67 |-751.83|15.07 |-0.05|-0.11|
|$w_5^*$      |-503783.21|171.74 |3.95  |0.56 |0.01 |
|$w_6^*$      |818011.41 |604.53 |-5.68 |0.79 |0.09 |
|$w_7^*$      |-785013.15|-111.73|-8.42 |0.83 |0.17 |
|$w_8^*$      |410395.73 |-692.14|-3.61 |0.78 |0.21 |
|$w_9^*$      |-90055.14 |401.63 |7.65  |0.69 |0.25 |

<!-- grade_id: cell-5156541a32d37722, points: 0, grade: false, solution: true -->
```python
def polynomial_curve_fit(D, M, lmb = 0):
    ''' This functions computes the value of a polynomial with weights w on data points x.
    
    Parameters
    ----------
    D : array
        Input dataset D.
    M : int
        The degree of the polynomial.
    lmb : float, optional
        Regularization parameter for polynomial curve fitting.
    
    Returns
    -------
    float
        Fitted weight vector w that minimizes the sum-of-squares function.
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7a30d44592cf4486, points: 5, grade: true, solution: false -->
```python
"""
Hidden test for polynomial_curve_fit.
"""
```

<!-- grade_id: cell-25af5f44d49bab7b, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 1.6
The polynomial curve fitting procedure can be extended to the case of multidimensional inputs. Assuming an input vector  of dimension $D$, namely $\mathbf{x} = (x_1, x_2, \dots, x_D)$, we can write the regression function $y$ as:
\begin{equation}
y(\mathbf{x}; \mathbf{w}) = \sum_{j = 0}^M \left( \sum_{n_1 + n_2 + ... + n_D = j} w_{n_1 n_2 ... n_D} x_1^{n_1} x_2^{n_2} ... x_D^{n_D} \right) \label{eqn:polynomial_multidimensional} \tag{5}
\end{equation}

In the last expression, $j$ refers to the order of the polynomial terms. The inner sum is over all the combinations of non-negative integers $n_1, n_2, \dots, n_D$, such that the constraint $n_1 + n_2 + \dots + n_D = j$ holds. The terms $n_1, n_2, \dots, n_D$ correspond to the exponent for each variable $x_1, x_2, \dots, x_D$ in their respective polynomial term.

Note that if $D = 1$, the above expression simplifies to the formula in equation \eqref{yxw}. The reason the second sum disappears is that there is only one combination of the non-negative integer $n_1$ for which the constraint $n_1 = j$ holds, which means that there is only  a single term to sum over.

Fitting the polynomial curve to a multidimensional input vector works analogously to the one-dimensional case. However, the number of parameters (the size of $\mathbf{w}$) becomes much larger, even when $D = 2$. Write down the general polynomial curve equation in \eqref{eqn:polynomial_multidimensional} for $D = 2$. How many parameters are needed in the two-dimensional case? Compare this to the number of parameters in the one-dimensional case.

<!-- grade_id: cell-af7fb0f41aed699d, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Exercise 2 (weight 25)
In this exercise, we consider the gradient descent algorithm for function minimization. When the function to be minimized is $E(\mathbf{x})$, the gradient descent iteration is  
\begin{equation*}
\mathbf{x}_{n+1} = \mathbf{x}_n - \eta \nabla E(\mathbf{x}_n) \tag{6}
\end{equation*}
where $\eta>0$ is the so-called learning-rate. In the following, we will apply gradient descent to the function
\begin{equation*}
h(x,y) = 100(y - x^2)^2 +(1 - x)^2 \label{banana} \tag{7}
\end{equation*}
### Exercise 2.1
Make a plot of the function $h$ over the interval $[-2 \leq x \leq 2] \times [-1 \leq y \leq 3]$. (Tip: Use the `plot_surface` function.) Can you guess from the plot if numerical minimization with gradient descent will be fast or slow for this function?

<!-- grade_id: cell-5dfbbae25832f037, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-d2b407ddc3fbb9fc, points: 0, grade: false, solution: true -->
```python
"""
Create function h.
"""
def h(x,y):
    # YOUR CODE HERE
    raise NotImplementedError()

"""
Declare x and y.
"""    
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e9cba008bca6b37a, points: 2, grade: true, solution: false -->
```python
"""
Hiddent test for function h.
"""
```

<!-- grade_id: cell-85ead8c8bd53e79c, points: 2, grade: true, solution: true -->
```python
"""
Create a function to plot h.
"""
def plot_h(x,y):
    # YOUR CODE HERE
    raise NotImplementedError()

plot_h(x,y)
```

### Exercise 2.2
Knowing that a critical point of a function is a point where the gradient vanishes, show that $(1, 1)$ is the unique critical point of $h$.  Prove that this point is a minimum for $h$. 

<!-- grade_id: cell-1336cb1e9d7d38bd, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 2.3
Write down the gradient descent iteration rule for this function. 

<!-- grade_id: cell-91163736650b6698, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 2.4
Implement gradient descent. Try some different values of $\eta$. Does the algorithm converge? How fast? Make plots of the trajectories on top of a contour plot of $h$. (Hint: have a look at the example contour_example.py on Brightspace for inspiration to plot contours of functions and trajectories). Report your findings. Explain why numerical minimization with gradient descent is slow for this function.

First implement the derivative of $h(x,y)$.

<!-- grade_id: cell-ebfab8f71f837375, points: 0, grade: false, solution: true -->
```python
def dh_dxy(x, y):
    """
    This function is the derivative of the function h(x, y).
    
    Parameters
    ----------
    x : float
        data point from x-axis
    y : float
        data point from y-axis
    
    Returns
    -------
    vals : array
        NumPy array of parameter values computed during minimization
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-14c1d4dd062d9e68, points: 1, grade: true, solution: false -->
```python
"""
Test for dh_dxy.
"""
assert np.array_equal(dh_dxy(1, 1), np.array([0, 0]))
```

Now implement the gradient descent algorithm.

<!-- grade_id: cell-e5e7a55f7e7d740f, points: 0, grade: false, solution: true -->
```python
def grad_descent(grad, val_init, eta, max_iter, tol):
    """ This function implements the gradient descent algorithm.
    
    Parameters
    ----------
    grad : function
        Returns the derivative of the function with respect to the pair (x, y).
    val_init : tuple
        Initial values for parameters
    eta : float
        Gradient descent learning rate
    max_iter : int
        Maximum number of gradient descent iterations
    tol : float
        Tolerance for detecting convergence
    
    Returns
    -------
    vals : array
        NumPy array of parameter values computed during minimization
    dists : array
        NumPy array of distances from the current point to the previous point
    tot_iter : int
        Number of performed gradient descent iterations
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9cce6b818f4aa52e, points: 3, grade: true, solution: false -->
```python
"""
Hidden test for grad_descent.
"""
```

Finally, run the gradient descent algorithm with different values of $\eta$.

<!-- grade_id: cell-76876feddeb4a5f0, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Explain what you see!

<!-- grade_id: cell-40c451de6dac1163, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Exercise 3 (weight 25)
Suppose we have two healthy but curiously mixed boxes of fruit, with one box containing 8 apples and 4 grapefruit and the other containing 15 apples and 3 grapefruit. One of the boxes is selected at random and a piece of fruit is picked (but not eaten) from the chosen box, with equal probability for each item in the box. The piece of fruit is returned and then once again from the *same* box a second piece is chosen at random. This is known as sampling with replacement. Model the box by random variable $B$, the first piece of fruit by variable $F_1$, and the second piece by $F_2$.
### Exercise 3.1
What is the probability that the first piece of fruit is an apple given that the second piece of fruit was a grapefruit? How can the result of the second pick affect the probability of the first pick?

<!-- grade_id: cell-20711d5530b05594, points: 9, grade: true, solution: true -->
YOUR ANSWER HERE

Please add the final result you got in the cell below! (Add it as a fraction, not an estimate. For example, write __1/3__, do not round to a number of decimals.)

<!-- grade_id: cell-6d4096e0b9231c28, points: 0, grade: false, solution: true -->
```python
"""
The variable p is probability of the first piece of fruit being
an apple given that the second piece of fruit was a grapefruit.
"""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-13f94c0b98af50e6, points: 0, grade: true, solution: false -->
```python
"""
Hidden check for value of variable p.
"""
```

### Exercise 3.2
Imagine now that after we remove a piece of fruit, it is not returned to the box. This is known as sampling without replacement. In this situation, recompute the probability that the first piece of fruit is an apple given that the second piece of fruit was a grapefruit. Explain the difference.

<!-- grade_id: cell-c5cabea66205f678, points: 9, grade: true, solution: true -->
YOUR ANSWER HERE

Please add the final result you got in the cell below! (Add it as a fraction, not an estimate. For example, write __1/3__, do not round to a number of decimals.)

<!-- grade_id: cell-bbf669c7d7e989eb, points: 0, grade: false, solution: true -->
```python
"""
The variable p is probability of the first piece of fruit being
an apple given that the second piece of fruit was a grapefruit
when the sampling was done without replacement.
"""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1aa31b8a1047f19a, points: 0, grade: true, solution: false -->
```python
"""
Hidden check for value of variable p.
"""
```

### Exercise 3.3
Starting from the initial situation (i.e., sampling with replacement), we add a dozen oranges to the first box and repeat the experiment. Show that now the outcome of the first pick has no impact on the probability that the second pick is a grapefruit. Are the two picks now dependent or independent? Explain your answer.

<!-- grade_id: cell-03574d2eb9eadf55, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

## Exercise 4 - Bonus (weight 10)
Given a joint probability density function over the random vector $\mathbf{X} = (X_1, X_2, X_3, X_4)$ that factorizes as
$$p(x_1, x_2, x_3, x_4) = p(x_1, x_4 | x_2) p(x_2, x_3 | x_1),$$
show (using the sum and product rules for marginals and conditionals) that the following independence statements hold, in which the symbol $\bot$ stands for (conditional) independence:
1. $ X_1 \bot X_2;$
2. $ X_3 \bot X_4 \,|\, X_1, X_2.$

<!-- grade_id: cell-fea97ebe52a73eb3, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE