## Students
Please fill in your names and S/U-numbers:
* Student 1 name, S/U-number:
* Student 2 name, S/U-number:
* Student 3 name, S/U-number:

# Statistical Machine Learning 2020
# Assignment 3
# Deadline: 2 December 2020
## Instructions
* You can __work in groups__ (= max 3 people). __Write the full name and S/U-number of all team members in the header above.__
* Make sure you __fill in any place that says__ `YOUR CODE HERE` or "YOUR ANSWER HERE" __including comments, derivations, explanations, graphs, etc.__ This means that the elements and/or intermediate steps required to derive the answer have to be in the report. (Answers like 'No' or 'x=27.2' by themselves are not sufficient, even when they are the result of running your code.) If an exercise requires coding, explain briefly what the code does (in comments). All figures should have titles (descriptions), axis labels, and legends (if applicable).
* Please do not add new cells unless necessary, try to write the answers only in the provided cells. Before you turn this problem in, __make sure everything runs as expected__. First, *restart the kernel* (in the menubar, select Kernel$\rightarrow$Restart) and then *run all cells* (in the menubar, select Cell$\rightarrow$Run All). The assignment was written in (and we strongly recommend using) Python 3 by using the corresponding Python 3 kernel for Jupyter.
* The assignment includes certain cells that contain tests. Most of the tests are marked as *hidden* and are used for automatic grading. NB: These hidden tests do not provide any feedback! There are also a couple of tests / checks that are visible, which are meant to help you avoid basic coding errors.
* __Upload the exercises to Brightspace as a single .zip file containing the submitter's S/U-number: 'SML20_as03_&lt;S/U-number&gt;.zip'__, for example 'SML20_as03_S123456.zip'. For those working in groups, it is sufficient if one team member uploads the solutions.
* For any problems or questions, send us an email, or just ask. Email addresses: G.Bucur@cs.ru.nl, Yuliya.Shapovalova@ru.nl, and tomc@cs.ru.nl.

## Introduction
Assignment 3 consists of:
1. The faulty lighthouse (40 points);
2. Neural network regression (30 points);
3. __Bayesian polynomial regression (30 points)__.

## Libraries

First, we import the basic libraries necessary to develop this assignment. Of course you are free to import further libraries, if required, in the allotted cells.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."

import numpy as np
import matplotlib.pyplot as plt

# Set fixed random seed for reproducibility
np.random.seed(2020)
```

## Bayesian polynomial regression (30 points)
In this exercise, we will consider _Bayesian_ treatment of polynomial regression. Recall polynomial regression 
\begin{equation}
t_{n} = \omega_{0} + \omega_{1}x_{n} + \omega_{2} x_{x}^{2}+\dots+\omega_{M}x_{n}^{M}+\epsilon_{n}, 
\end{equation}
where $\epsilon_{n}\sim N(0, \sigma^{2})$. In the vector form we have 
\begin{equation}
t_{n} = \boldsymbol{w}^{T}\boldsymbol{x}_{n}+\epsilon_{n}, 
\end{equation}
where $\boldsymbol{w}=[w_{0}, \dots, w_{K}]^{T}$ and $\boldsymbol{x}_{n}=[1,x_{n}, x_{n}^{2}, \dots, x_{n}^{K}]$. Further, let us stack all responses in one vector $\boldsymbol{t}=[t_{1}, \dots, t_{N}]^{T}$, all inputs in a single matrix $\boldsymbol{X}=[\boldsymbol{x}_{1}, \boldsymbol{x}_{2}, \dots, \boldsymbol{x}_{N}]^{T}$. Then we get for the whole data set
\begin{equation}
\boldsymbol{t}=\boldsymbol{X}\boldsymbol{w}+\boldsymbol{\epsilon},
\end{equation}
where $\boldsymbol{\epsilon}=[\epsilon_{1}, \dots, \epsilon_{N}]^{T}$. 
Assume that we know the true value of  $\sigma^{2}$. 

1. Derive the posterior distribution of $\boldsymbol{\omega}$, i.e. $p(\boldsymbol{\omega}|\boldsymbol{t}, \boldsymbol{X}, \sigma^{2})$.  
Hint: Use the prior $p(\boldsymbol{\omega}|\boldsymbol{\mu_{0}}, \boldsymbol{\Sigma_{0}})=N(\boldsymbol{\mu_{0}}, \boldsymbol{\Sigma_{0}})$ and the fact that the posterior should be Gaussian. 

<!-- grade_id: cell-d12dbe36ec039c0b, points: 8, grade: true, solution: true -->
YOUR ANSWER HERE

2. Take the first order polynomial, the the inputs are $\boldsymbol{x_{n}}=[1, x_{n}]^{T}$. Let $\boldsymbol{\mu_{0}}=[0,0,\dots,0]^{T}$. Then the posterior mean for the linear Gaussian model is
\begin{equation}
\boldsymbol{\mu_{\omega}}=(\boldsymbol{X}^{T}\boldsymbol{X}+ \sigma^2 \boldsymbol{\Sigma_{0}^{-1}})^{-1}\boldsymbol{X}^{T}\boldsymbol{t}.
\label{eq:MAP} \tag{1}
\end{equation}
Recall also the regularised least squares solution:
\begin{equation}
\hat{\boldsymbol{\omega}}=(\boldsymbol{X}^{T}\boldsymbol{X}+N\lambda\boldsymbol{I})^{-1}\boldsymbol{X}^{T}\boldsymbol{t}.
\label{eq:MLE} \tag{2}
\end{equation}
Find $\boldsymbol{\Sigma_{0}}$ that makes Equation \eqref{eq:MAP} and Equation \eqref{eq:MLE} identical. Reflect on the similarity between MAP solution and regularized least squares. Comment on what it implies for the effect of the prior.

<!-- grade_id: cell-cb4a19df380490ad, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

3. Generate data from the polynomial function
\begin{equation}
t= 5x^{3} - x^{2} +x
\end{equation}
and add Gaussian noise with the mean zero and variance 150 and plot your data.

First define the function $t$.

<!-- grade_id: cell-ece70abe254c227c, points: 0, grade: false, solution: true -->
```python
def f(x):
    """
    Define the polynomial function f(x).
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5db410b0084695b0, points: 1, grade: true, solution: false -->
```python
"""
Test the polynomial function f(x).
"""
```

Now plot the data generated by the function with a mean of 0 and a variance of 150.

<!-- grade_id: cell-98a8f3ba37e9002f, points: 3, grade: true, solution: true -->
```python
"""
Plot the polynomial function t(x) with added Gaussian noise.
"""
# YOUR CODE HERE
raise NotImplementedError()
```

4. Compute the marginal likelihood for the models from first to seventh order and make a plot of the marginal likelihood for these models (polynomial order on the x-axis and marginal likelihood value on the y-axis). Use a Gaussian prior on $\boldsymbol{\omega}$ with **zero mean** and an **identity covariance matrix**.

Hints: 
1) If you are computing log marginal likelihood (which you probably are), you might want to take exponent of that when making the plots. You can use `matplotlib.pyplot.bar` for making the bar plot.  
2) The marginal likelihood (also known as the model evidence) for our Gaussian model is defined as $p(\boldsymbol{t}|\boldsymbol{X}, \boldsymbol{\mu_{0}}, \boldsymbol{\Sigma_{0}})$. Using Gaussian distributions as before we can compute
\begin{equation}
p(\boldsymbol{t}|\boldsymbol{X}, \boldsymbol{\mu_{0}}, \boldsymbol{\Sigma_{0}})= N(\boldsymbol{X}\boldsymbol{\mu_{0}}, \sigma^{2}\boldsymbol{I}_{N}+\boldsymbol{X}\boldsymbol{\Sigma_{0}}\boldsymbol{X}^{T}).
\end{equation}

<!-- grade_id: cell-cadd5890106dd435, points: 0, grade: false, solution: true -->
```python
def marginal_loglikelihood(X, N, t, sigma=1):
    """
    Calculate the marginal likelihood.

    Parameters
    ----------
    X : Nx1 array
        The array containing the random data points.
    N : numeric
        The number of data points.
    t : Nx1 array
        The array containing the random data points with Gaussian noise.

    Returns
    -------
    list
        Marginal loglikelihoods for different polynomial orders.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d78cc77a3b106231, points: 4, grade: true, solution: false -->
```python
"""
Test for marginal_loglikelihood.
"""
assert type(marginal_loglikelihood(np.full((1, 1), 0), 1, np.full((1, 1), 0))) is list
```

Now plot the results with the help of `marginal_loglikelihood`.

<!-- grade_id: cell-8ecce257dda193fb, points: 2, grade: true, solution: true -->
```python
"""
Plot the results of the marginal loglikelihood.
"""
# YOUR CODE HERE
raise NotImplementedError()
```

Comment on your results. Which model would you choose according to the marginal likelihood?  

<!-- grade_id: cell-3f01a1d2ee023123, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

4. How would the prior affect the choice of polynomial using marginal likelihood? Let $\boldsymbol{\Sigma_{0}}=\sigma_{0}^{2}\boldsymbol{I}$ and vary $\sigma_{0}^{2}$: what happens when you increase and decrease $\sigma_{0}^{2}$? Plot the marginal likelihood fror various polynomial orders (from 1 to 7) for $\sigma_{0}^{2}=0.1$ , $\sigma_{0}^{2}=0.3$, $\sigma_{0}^{2}=0.4$, $\sigma_{0}^{2}=0.7$, $\sigma_{0}^{2}=1.3$, $\sigma_{0}^{2}=1.4$, $\sigma_{0}^{2}=1.7$.

<!-- grade_id: cell-5a8f27b42908767d, points: 5, grade: true, solution: true -->
```python
"""
Plot the results of the marginal loglikelihood with different sigmas.
"""
# YOUR CODE HERE
raise NotImplementedError()
```

Comment on the effect of changing $\sigma_{0}^{2}$. What does it in general imply for such type of modeling choice (i.e. for Bayesian way of estimation for the polynomial regression)?

<!-- grade_id: cell-e2e8c5e272656373, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE